# Structs

In this notebook we'll explore **structs**, which is a key feature of Julia.

Multiple dispatch makes software *generic* and *fast*!

#### Starting with the familiar

```julia
struct myObj 
    field1
    field2
end

myobj1 = myObj("Hello","World")
```

Output:
```
myObj("Hello", "World")
```

```julia
myobj1.field1
```

Output:
```
"Hello"
```

```julia
myobj1.field1 = "Hi"
```

Output:
```
@ Base ./Base.jl:41
 [2] top-level scope
   @ In[4]:1
```

```julia
mutable struct Person
    name::String 
    age::Float64
end
```

```julia
Bryan = Person("Bryan", 30)
```

Output:
```
Person("Bryan", 30.0)
```

```julia
Bryan.age+=1
```

Output:
```
31.0
```

```julia
dude = Person(20,30)
```

Output:
```
@ Main ./In[5]:2
 [2] top-level scope
   @ In[9]:1
```

```julia
mutable struct Person2
    name::String 
    age::Float64
    isActive

    function Person2(name,age)
        new(name,age,true)
    end
end
```

```julia
newperson = Person2("Kelly", 30)
```

Output:
```
Person2("Kelly", 30.0, true)
```

```julia
function birthday(person::Person2)
    person.age +=1
end
birthday(newperson)
```

Output:
```
31.0
```

```julia
newperson.age
```

Output:
```
31.0
```

Let's see what happens when we use `@which` with the addition operator!

```julia
@which 3.0 + 3.0
```

Output:
```
+([90mx[39m::[1mT[22m, [90my[39m::[1mT[22m) where T<:Union{Float16, Float32, Float64}
[90m     @[39m [90mBase[39m [90m[4mfloat.jl:409[24m[39m
```

And we can continue to add other methods to our generic function `foo`. Let's add one that takes the ***abstract type*** `Number`, which includes subtypes such as `Int`, `Float64`, and other objects you would think of as numbers:

```julia
foo(x::Number, y::Number) = println("My inputs x and y are both numbers!")
```

Output:
```
foo (generic function with 3 methods)
```

This method for `foo` will work on, for example, floating point numbers:

```julia
foo(3.0, 4.0)
```

Output:
```
My inputs x and y are both numbers!
```

We can also add a fallback, duck-typed method for `foo` that takes inputs of any type:

```julia
foo(x, y) = println("I accept inputs of any type!")
```

Output:
```
foo (generic function with 4 methods)
```

Given the methods we've already written for `foo` so far, this method will be called whenever we pass non-numbers to `foo`:

```julia
v = rand(3)
foo(v, v)
```

Output:
```
I accept inputs of any type!
```

### Exercises

#### 9.1

Extend the function `foo`, adding a method that takes only one input argument, which is of type `Bool`, and returns "foo with one boolean!"

```julia
foo(x::Bool) = "foo with one boolean!"
foo(true)
```

Output:
```
"foo with one boolean!"
```

#### 9.2

Check that the method being dispatched when you execute 
```julia
foo(true)
```
is the one you wrote.

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert foo(true) == "foo with one boolean!"
```